---
id: 587d78b2367417b2b2512b0e
title: Add Items to an Array with push() and unshift()
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

An array's length, like the data types it can contain, is not fixed. Arrays can be defined with a length of any number of elements, and elements can be added or removed over time; in other words, arrays are <dfn>mutable</dfn>. In this challenge, we will look at two methods with which we can programmatically modify an array: `Array.push()` and `Array.unshift()`.

**Real-world analogy:** Think of an array like a **train**. Each item in the array is a train car.
*   **`.push()`** is like attaching a new car to the **very end** of the train.
*   **`.unshift()`** is like adding a new car to the **very front**, right behind the engine.

Both methods take one or more elements as parameters and add those elements to the array.

**Step-by-step breakdown:**

```js
let romanNumerals = ['XXI', 'XXII'];

// 1. Adding to the front (unshift)
romanNumerals.unshift('XIX', 'XX'); 
// ['XIX', 'XX', 'XXI', 'XXII']

// 2. Adding to the end (push)
let twentyThree = 'XXIII';
romanNumerals.push(twentyThree);
// ['XIX', 'XX', 'XXI', 'XXII', 'XXIII']
```

Notice that we can also pass variables, which allows us even greater flexibility in dynamically modifying our array's data.

**Simple real-world use case:**

Imagine you are building a **music playlist**:
- Use `.push()` to add a new song to the end of the queue.
- Use `.unshift()` to "play next" by adding a song to the very top of the list.

```js
let playlist = ["Song A", "Song B"];
playlist.push("Song C");      // Adds to the end
playlist.unshift("Song Z");   // Puts it at the top
```

# --instructions--

We have defined a function, `mixedNumbers`, which we are passing an array as an argument. Modify the function by using `push()` and `unshift()` to add `'I', 2, 'three'` to the beginning of the array and `7, 'VIII', 9` to the end so that the returned array contains representations of the numbers 1-9 in order.

# --hints--

`mixedNumbers(["IV", 5, "six"])` should now return `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]`

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

The `mixedNumbers` function should utilize the `push()` method

```js
assert(mixedNumbers.toString().match(/\.push/));
```

The `mixedNumbers` function should utilize the `unshift()` method

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
