---
id: 56533eb9ac21ba0edf2244d4
title: Comparison with the Greater Than Operator
challengeType: 1
forumTopicId: 16786
dashedName: comparison-with-the-greater-than-operator
---

# --description--

The greater than operator (`>`) compares the values of two numbers. If the number to the left is greater than the number to the right, it returns `true`. Otherwise, it returns `false`.

**Real-world analogy:** Think of checking if someone is tall enough to ride a roller coaster. If their height is greater than the minimum requirement, they can ride (`true`). Otherwise, they can't (`false`).

**Step-by-step breakdown:**

```js
5  > 3   // Is 5 greater than 3?  YES → true
2  > 3   // Is 2 greater than 3?  NO  → false
```

**Type coercion note:** Like the equality operator, the greater than operator will convert data types of values while comparing. So if you compare a number to a number-like string, JavaScript converts the string to a number first:

```js
7   > '3'  // '3' becomes 3 → Is 7 > 3? YES → true
'1' >  9   // '1' becomes 1 → Is 1 > 9? NO  → false
```

**More examples:**

```js
5   >  3  // true
7   > '3' // true
2   >  3  // false
'1' >  9  // false
```

**Simple real-world use case:**

An age check on a website — only let users in if they're over 18:

```js
let userAge = 20;
if (userAge > 18) {
  console.log("Welcome!");
} else {
  console.log("Sorry, you must be over 18.");
}
```

# --instructions--

Add the greater than operator to the indicated lines so that the return statements make sense.

# --hints--

`testGreaterThan(0)` should return the string `10 or Under`

```js
assert(testGreaterThan(0) === '10 or Under');
```

`testGreaterThan(10)` should return the string `10 or Under`

```js
assert(testGreaterThan(10) === '10 or Under');
```

`testGreaterThan(11)` should return the string `Over 10`

```js
assert(testGreaterThan(11) === 'Over 10');
```

`testGreaterThan(99)` should return the string `Over 10`

```js
assert(testGreaterThan(99) === 'Over 10');
```

`testGreaterThan(100)` should return the string `Over 10`

```js
assert(testGreaterThan(100) === 'Over 10');
```

`testGreaterThan(101)` should return the string `Over 100`

```js
assert(testGreaterThan(101) === 'Over 100');
```

`testGreaterThan(150)` should return the string `Over 100`

```js
assert(testGreaterThan(150) === 'Over 100');
```

You should use the `>` operator at least twice

```js
assert(__helpers.removeJSComments(code).match(/val\s*>\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterThan(val) {
  if (val) {  // Change this line
    return "Over 100";
  }

  if (val) {  // Change this line
    return "Over 10";
  }

  return "10 or Under";
}

testGreaterThan(10);
```

# --solutions--

```js
function testGreaterThan(val) {
  if (val > 100) {  // Change this line
    return "Over 100";
  }
  if (val > 10) {  // Change this line
    return "Over 10";
  }
  return "10 or Under";
}
```
