---
id: 56533eb9ac21ba0edf2244c4
title: Return Early Pattern for Functions
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

When a `return` statement is reached, the execution of the current function stops and control returns to the calling location.

**Real-world analogy:** Imagine a security guard at a club entrance. If someone doesn't meet the age requirement, the guard says "Sorry, you can't come in" and turns them away — they don't even step inside. The `return` statement works the same way: it exits the function immediately, skipping everything that comes after it.

**Step-by-step breakdown:**

```js
function myFun() {
  console.log("Hello");   // ✅ This runs — it comes before return
  return "World";          // 🛑 Function stops here and sends back "World"
  console.log("byebye");  // ❌ This NEVER runs — it comes after return
}
myFun();
```

The above will display the string `Hello` in the console, and return the string `World`. The string `byebye` will never display in the console, because the function exits at the `return` statement.

**Simple real-world use case:**

The early return pattern is often used to **validate input** before doing any real work. For example, checking that a number is positive before doing a calculation:

```js
function calculateSquareRoot(num) {
  if (num < 0) {
    return "Error: Cannot take square root of a negative number";
  }
  // Only reaches here if num is 0 or positive
  return Math.sqrt(num);
}
```

This keeps your code clean — handle the "bad" cases first, then focus on the normal logic.

# --instructions--

Modify the function `abTest` so that if `a` or `b` are less than `0` the function will immediately exit with a value of `undefined`.

**Hint**  
Remember that <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> is a keyword</a>, not a string.

# --hints--

`abTest(2, 2)` should return a number

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` should return `8`

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` should return `undefined`

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` should return `undefined`

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` should return `18`

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` should return `12`

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` should return `0`

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
