---
id: 56533eb9ac21ba0edf2244bc
title: Shopping List
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Create a shopping list in the variable `myList`. The list should be a multi-dimensional array containing several sub-arrays.

**What is a multi-dimensional array?**

A multi-dimensional array is simply an array that contains other arrays inside it — like a list of lists. Think of it like a **grocery list** where each item has a name and a quantity:

| Item | Quantity |
|------|----------|
| Eggs | 12 |
| Milk | 2 |
| Bread | 1 |

In JavaScript, you represent each row as a sub-array — a small array inside the bigger array:

```js
["Eggs", 12]    // name first, then quantity
["Milk", 2]
["Bread", 1]
```

Then the whole list becomes an array of these sub-arrays:

```js
const groceryList = [
  ["Eggs", 12],
  ["Milk", 2],
  ["Bread", 1]
];
```

**Step-by-step breakdown of the format:**

The first element in each sub-array should contain a string with the name of the item. The second element should be a number representing the quantity, for example:

```js
["Chocolate Bar", 15]
```

There should be at least 5 sub-arrays in the list.

**Why is this useful?**

This pattern appears everywhere in real apps — from e-commerce product lists to seating charts and schedule grids. Grouping related data (name + quantity) together in a sub-array keeps things organized and easy to loop through.

# --hints--

`myList` should be an array.

```js
assert(isArray);
```

The first elements in each of your sub-arrays should all be strings.

```js
assert(hasString);
```

The second elements in each of your sub-arrays should all be numbers.

```js
assert(hasNumber);
```

You should have at least 5 items in your list.

```js
assert(count > 4);
```

# --seed--

## --after-user-code--

```js
var count = 0;
var isArray = false;
var hasString = false;
var hasNumber = false;
(function(list){
  if(Array.isArray(myList)) {
    isArray = true;
    if(myList.length > 0) {
      hasString = true;
      hasNumber = true;
      for (var elem of myList) {
        if(!elem || !elem[0] || typeof elem[0] !== 'string') {
          hasString = false;
        }
        if(!elem || typeof elem[1] !== 'number') {
          hasNumber = false;
        }
      }
    }
    count = myList.length;
    return JSON.stringify(myList);
  } else {
    return "myList is not an array";
  }

})(myList);
```

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
