---
id: 56533eb9ac21ba0edf2244a8
title: Storing Values with the Assignment Operator
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

In JavaScript, you can store a value in a variable with the <dfn>assignment</dfn> operator (`=`).

**Real-world analogy:** Think of a variable like a labeled box. The variable name is the label on the box, and the value is whatever you put inside it. The `=` operator is the action of placing something into the box.

```js
myVariable = 5;
```

This assigns the `Number` value `5` to `myVariable`.

**Step-by-step breakdown:**

```js
var myVar;   // Step 1: Create an empty box labeled "myVar"
myVar = 5;   // Step 2: Put the value 5 inside the box
```

After step 2, whenever `myVar` appears in the code, JavaScript treats it as if it is `5`.

**Simple use case:** Imagine you're building a shopping cart. You could store the item price in a variable, and refer to it whenever you need to calculate the total:

```js
var itemPrice;
itemPrice = 9.99;
// Now you can use itemPrice anywhere in your code instead of repeating 9.99
```

If calculations appear to the right of `=`, those are performed **first**, and then the result is stored in the variable on the left.

# --instructions--

Assign the value `7` to variable `a`.

# --hints--

You should not change code above the specified comment.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` should have a value of 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
